# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    
    # 🧸💬 Regular expression, name of person, place, vehicle and objects are 
    # express by themself. They can be subject of the sentence and identify
    # among group of the same or different purpose.
    expression = r'([A-Z][a-z]\w+)'
    _temp = re.findall(expression, simple_string)
    
    return _temp

print( names() )
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()
        
        # 🧸💬 Name and commons are identical and more indent with behaviour.
        expression = r'([A-Z]\w+).([A-Z][a-z]\w+)(: B)'
        _temp = re.findall(expression, grades)
        
        result = [str(x) + " " + str(y) for (x,y,z) in _temp]

    return result

# _temp = grades()
# print( _temp )
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    
    i_count = 0
    ls_hostname = []
    ls_username = []
    ls_time = []
    ls_request = []
    all_text = ""
    result = []
    
    with open("assets/logdata.txt", "r") as file:
        i_count = i_count + 1;
        logdata = file.read()
        all_text = logdata.split("\n")
    
    for logdata in all_text:
        
        if len(logdata) == 0:
            break;
        
        # 🧸💬 IP-address compose of number and '.' that preriodic significants.
        # hostname_exp = r'(\w+)[.](\w+)[.](\w+)[.]([0-9]\w+)'
        hostname_exp = r'(\w+[/.])(\w+[/.])(\w+[/.])(\w+[ ])'
        _temp = re.findall(hostname_exp, logdata)

        hostname = ""
        if len(_temp) > 0:
            hostname = f"{ _temp[0][0] }{ _temp[0][1] }{ _temp[0][2] }{ _temp[0][3] }"
        
        # 🧸💬 Username compose of common and capital, they are significant self and friends.
        username_exp = r'([-][ ])([a-zA-Z0-9]\w+)'
        _temp = re.findall(username_exp, logdata)
        username = ""
        if len(_temp) > 0:
            username = f"{ _temp[0][1] }"
        
        # 🧸💬 Date and time, significant and can be starter.
        ##############################################
        date_exp = r'[\[]+(\w+)'
        _temp = re.search(date_exp, logdata) 
        
        date = logdata[_temp.span()[0]+1:_temp.span()[1]]
        
        _tempstring = logdata[_temp.span()[1] + 1:]
        month_exp = r'(\w+)[/]'
        _temp = re.search(month_exp, _tempstring)
        
        month = _tempstring[_temp.span()[0]:_temp.span()[1] - 1]
        
        _tempstring = _tempstring[_temp.span()[1]:]
        year_exp = r'(\w+)[:]'
        _temp = re.search(year_exp, _tempstring)
        
        year = _tempstring[_temp.span()[0]:_temp.span()[1] - 1]
        
        timezone_exp = r'([ -]\w+)'
        _temp = re.findall(timezone_exp, _tempstring)
        timezone = _temp[0].strip()
        ##############################################
        
        # 🧸💬 Time a good interval for tasks.
        time_exp = r'([:])(\w+)'        
        _temp = re.findall(time_exp, logdata)   

        hour, minute, time = 0, 0, 0
        if len(_temp) > 0:
            hour = _temp[0][1]
            minute = _temp[1][1]
            sec = _temp[2][1]
            time = f"{hour}:{minute}:{sec}"

        # 🧸💬 Request common message we need to maintain and continue.
        # request_exp = r'([ -./\w+]\w+)'
        request_exp = r'(["\%\+ \.\-\/\\""]\w+)'
        _temp = re.findall(request_exp, logdata)  
        
        idx_HTTPmethod = 0;
        bStartconcat = False;
        for item in _temp[:-2]:
            item = item.replace("\"", "")
            if item in ("POST", "GET", "PUT", "HEAD", "DELETE", "PATCH"):
                item = item + " ";
                bStartconcat = True;
                
            if bStartconcat:
                request = f"{ request }{ item }";
            else:
                request = "";

        temp = f'{date}/{month}/{year}:{ str(time).strip() } {timezone}';
        
        message = {}
        if len(hostname.strip()) > 0:
            message["host"] = hostname.strip();
        else:
            message["host"] = ' '
        if len(username.strip()) > 0:
            message["user_name"] = username.strip();
        else:
            message["user_name"] = ' '
        message["time"] = temp;
        message["request"] = request.strip();
        
        result.append(message)
    
    return result
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python
#🧸💬 🐑💬 ➰ Yeahhhh~
```

```python

```